## Problem 1 - Simple temperature calculator (6 points)

In the first problem your aim is to create a function that converts the input temperature from degrees Fahrenheit to degrees Celsius. Follow closely the instructions and use **exactly** the same variable, function, and parameter names as stated in the instructions - otherwise the automatic tests will fail.

Include the following elements in your solution:

  - Create a function called `fahr_to_celsius` 

  - The function should have one input parameter called `temp_fahrenheit`

  - Inside the function, create a variable called `converted_temp` to which you should assign the conversion result (i.e., the input Fahrenheit temperature converted to Celsius)

  - Return the converted value from the function back to the user
  
  - Comment your code and add a docstring that explains how to use your `fahr_to_celsius` function (i.e., you should write the purpose of the function, parameters, and returned values)  
  
Here is the conversion formula from Fahrenheit to Celsius:

![](img/Fahrenheit_to_Celsius_formula.PNG)

  
If everything in your script is working properly the following test case should work:

  ```python
  >>> print("32 degrees Fahrenheit in Celsius is:", fahr_to_celsius(32))
  32 degrees Fahrenheit in Celsius is: 0.0
  ```

<!-- grade_id: fahr_to_celsius, points: 0, grade: false, solution: true -->
```python
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

**Use** your newly created function to print answers to the following questions:

```python
# 1. What is 48째 Fahrenheit in Celsius? 

# 2. What about 71째 Fahrenheit in Celsius?
```

### Problem 1 tests

Run the code cells below to check your answer. Note that you are not able to modify the test cells. Your solution is on the right track if you don't see any errors when running these code cells. If you receive an error, you need to change something in the code cells above.

<!-- grade_id: problem_1_fahr_to_celsius_test1, points: 2, grade: true, solution: false -->
```python
from nose.tools import ok_, assert_equal
import inspect

# Check that function exists
ok_(inspect.isfunction(fahr_to_celsius))
```

<!-- grade_id: problem_1_fahr_to_celsius_test2, points: 1, grade: true, solution: false -->
```python
# Check that the function has a single parameter and the parameter name is correct
params = list(inspect.signature(fahr_to_celsius).parameters.keys())
ok_(len(params)==1)
ok_(params[0]=='temp_fahrenheit')
```

<!-- grade_id: problem_1_fahr_to_celsius_test3, points: 1, grade: true, solution: false -->
```python
# Check that the function produces correct answers for:
# 1. What is 48째 Fahrenheit in Celsius? 
assert_equal(round(fahr_to_celsius(48), 2), 8.89)
```

<!-- grade_id: problem_1_fahr_to_celsius_test4, points: 1, grade: true, solution: false -->
```python
# 2. What about 71째 Fahrenheit in Celsius?
assert_equal(round(fahr_to_celsius(71), 2), 21.67)
```

### Check that your code is readable 

- Make sure you used the given variable names
- Check that you have added necessary comments to your code
- Check that your function has a docstring that describes what it does


### Questions

We would like you to think about and answer the following questions based on the materials and ideas that you learned during the lecture:

  1. Is the concept of function clear to you? If not, what do you not understand?
  2. What are some of the benefits of using functions?
  
Write your answers below (**activate the cell by double clicking**):

<!-- grade_id: problem_1_answers, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

#### Done!

That's it! Now you are ready to continue with [Problem 2](Exercise-4-problem-2.ipynb).