## Problem 3 -  Applying the classifier

Finally, we can bring together the pieces that we have created in problems 1 and 2. Here, your aim is to take
advantage of your new functions and reclassify a dataset of temperatures in Fahrenheit into four different classes.


### Part 1: Creating a script file

 - Create a new **script** -file called `temp_functions.py` into your `exercise-5` directory as we saw in Lesson 5 materials. Copy and paste the functions that you wrote in Problem 1 and Problem 2, into that script file (i.e. functions `fahr_to_celsius` and `temp_classifier`).
 - Add a docstring to the script file (in addition to separate docstrings in the functions).
 - Save the script file (CTRL + S), and continue working on **this** Notebook (i.e. `Exercise-5-problem-3.ipynb`).

### Part 2: Convert Fahrenheit temperatures to Celsius

 1. Store Fahrenheit temperatures into a list `temp_data` (just execute the existing code cell).

```python
# List of half-hourly temperature values (in degrees Fahrenheit) for one week
temp_data =  [19, 21, 21, 21, 23, 23, 23, 21, 19, 21, 19, 21, 23, 27, 27, 28, 30, 30, 32, 32, 32, 32, 
              34, 34, 34, 36, 36, 36, 36, 36, 36, 34, 34, 34, 34, 34, 34, 32, 30, 30, 30, 28, 28, 27,
              27, 27, 23, 23, 21, 21, 21, 19, 19, 19, 18, 18, 21, 27, 28, 30, 32, 34, 36, 37, 37, 37, 
              39, 39, 39, 39, 39, 39, 41, 41, 41, 41, 41, 39, 39, 37, 37, 36, 36, 34, 34, 32, 30, 30,
              28, 27, 27, 25, 23, 23, 21, 21, 19, 19, 19, 18, 18, 18, 21, 25, 27, 28, 34, 34, 41, 37, 
              37, 39, 39, 39, 39, 41, 41, 39, 39, 39, 39, 39, 41, 39, 39, 39, 37, 36, 34, 32, 28, 28,
              27, 25, 25, 25, 23, 23, 23, 23, 21, 21, 21, 21, 19, 21, 19, 21, 21, 19, 21, 27, 28, 32,
              36, 36, 37, 39, 39, 39, 39, 39, 41, 41, 41, 41, 41, 41, 41, 41, 41, 39, 37, 36, 36, 34,
              32, 30, 28, 28, 27, 27, 25, 25, 23, 23, 23, 21, 21, 21, 19, 19, 19, 19, 19, 19, 21, 23,
              23, 23, 25, 27, 30, 36, 37, 37, 39, 39, 41, 41, 41, 39, 39, 41, 43, 43, 43, 43, 43, 43,
              43, 43, 43, 39, 37, 37, 37, 36, 36, 36, 36, 34, 32, 32, 32, 32, 30, 30, 28, 28, 28, 27,
              27, 27, 27, 25, 27, 27, 27, 28, 28, 28, 30, 32, 32, 32, 34, 34, 36, 36, 36, 37, 37, 37,
              37, 37, 37, 37, 37, 37, 36, 34, 30, 30, 27, 27, 25, 25, 23, 21, 21, 21, 21, 19, 19, 19,
              19, 19, 18, 18, 18, 18, 18, 19, 23, 27, 30, 32, 32, 32, 32, 32, 32, 34, 34, 34, 34, 34,
              36, 36, 36, 36, 36, 32, 32, 32, 32, 32, 32, 32, 32, 30, 30, 30, 30, 30, 30, 30, 30, 30,
              30, 30, 30, 30, 28, 28]
```

**Part 2 continues**

 2. Import the `fahr_to_celsius` and `temp_classifier` functions from `temp_functions.py`
 3. Create an empty list called `temp_classes` (which will be filled with temperature class numbers later)
 4. **Convert the Fahrenheit temperatures** in the `temp_data` list into Celsius and **classify the Celsius temperatures** using the importend functions. Output should be stored in the `temp_classes` list:
 
 **You can do the temperature conversion and classification inside one for-loop:**
 
     - Iterate over the list of Fahrenheit temperatures (create a `for`-loop), and inside the loop:
        1. Create a new variable called `temp_celsius` in which you should assign the temperature in Celsius using the `fahr_to_celsius` function to convert the Fahrenheit temperature into Celsius.
        2. Create a new variable called `temp_class` in which you should assign the temperature class number (0, 1, 2, or 3) using the `temp_classifier` function
        3. Add the `temp_class` value to the `temp_classes` list

<!-- grade_id: functions_usage, points: 0, grade: false, solution: true -->
```python
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

**Part 2 continues**

6. Calculate how many temperatures there are in each temperature class:

    1. Create four variables called `zeros`, `ones`, `twos`, and `threes` 
    2. Count and assign to each variable how many times values 0, 1, 2, and 3 are present in the `temp_classes` list and print out the results below. I.e. assign the value of how many `0` values the list has into variable `zeros`, and so on. 
  
**TIP**: You might want to consider using a [**count()** -function](https://www.tutorialspoint.com/python3/list_count.htm) OR a for loop to handle this.

```python
# 1. How many 0 values exist in temp_classes -list?

# 2. How many 1 values exist in temp_classes -list?

# 3. How many 2 values exist in temp_classes -list?

# 4. How many 3 values exist in temp_classes -list?
```

### Problem 3 tests

Run the code cells below to check your answer. Note that you are not able to modify the test cells. Your solution is on the right track if you don't see any errors when running these code cells.

<!-- grade_id: problem_3_functions_usage_test1, points: 1, grade: true, solution: false -->
```python
from nose.tools import ok_, assert_equal
import inspect

# Check that functions are in the namespace
ok_(inspect.isfunction(fahr_to_celsius))
ok_(inspect.isfunction(temp_classifier))
```

<!-- grade_id: problem_3_functions_usage_test2, points: 1, grade: true, solution: false -->
```python
# Check that variable has been created
ok_('temp_celsius' in locals())
ok_('temp_class' in locals())
ok_('temp_classes' in locals())

# Check that temp_classes is a list
ok_(type(temp_classes) == list)
```

<!-- grade_id: problem_3_functions_usage_test3, points: 1, grade: true, solution: false -->
```python
# Check that the functions have a single parameter
t_params = list(inspect.signature(temp_classifier).parameters.keys())
f_params = list(inspect.signature(fahr_to_celsius).parameters.keys())
ok_(len(t_params)==1)
ok_(len(f_params)==1)
```

<!-- grade_id: problem_3_functions_usage_test4, points: 1, grade: true, solution: false -->
```python
# Check that the variable "zeros" exists and is correct
ok_('zeros' in locals())
```

<!-- grade_id: problem_3_functions_usage_test5, points: 1, grade: true, solution: false -->
```python
# Check that the variable "ones" exists and is correct
ok_('ones' in locals())
```

<!-- grade_id: problem_3_functions_usage_test6, points: 1, grade: true, solution: false -->
```python
# Check that the variable "twos" exists and is correct
ok_('twos' in locals())
```

<!-- grade_id: problem_3_functions_usage_test7, points: 1, grade: true, solution: false -->
```python
# Check that the variable "threes" exists and is correct
ok_('threes' in locals())
```

 ### Finishing up
 
 - Remember to check that your have commented your code and used the required variable names.
 - Make sure to commit and push also your `temp_functions.py` -script to your own GitHub repository for Exercise 5!
 - If you have any comments and concerns at this point, you can write them below:




<!-- grade_id: problem_3_manual_grade, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

#### All done!

That's it! Now you have finished all the problems for Exercise 5!